<!--NAVIGATION-->

<a href="https://colab.research.google.com/github/bpesquet/machine-learning-katas/blob/master/classic-datasets/Breast_Cancer.ipynb"><img align="left" src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open in Colab" title="Open in Google Colaboratory"></a>


# Kata: Breast Cancer Dataset

| Learning type | Activity type | Objective |
| - | - | - |
| Supervised | Binary classification | Predict if a tumor is benign or malignant |


## Instructions

This is a self-correcting exercise generated by [nbgrader](https://github.com/jupyter/nbgrader). 

Complete the cells beginning with `# YOUR CODE HERE` and run the subsequent cells to check your code.

## About the dataset

The [Breast Cancer][1] dataset is used for multivariate binary classification. There are 569 total samples with 30 features each. Features were computed from a digitized image of a fine needle aspirate of a breast mass. They describe characteristics of the cell nuclei present in the image.

![](images/breast-cancer-logo.jpg)

[1]: https://archive.ics.uci.edu/ml/datasets/Breast+Cancer+Wisconsin+(Diagnostic)

## Package setup

```python
import matplotlib
import matplotlib.pyplot as plt
import seaborn as sns
import pandas as pd
from sklearn.datasets import load_breast_cancer
from sklearn.linear_model import LogisticRegression
import keras.optimizers
from keras.models import Sequential
from keras.layers import Dense
from keras.utils import to_categorical
#from nose.tools import assert_equal, assert_true

# Display plots inline, change default figure size and change plot resolution to retina
%matplotlib inline
%config InlineBackend.figure_format = 'retina'
# Set Seaborn aesthetic parameters to defaults
sns.set()
```

```python
import numpy as np
from collections import Counter

class KNN:
    def __init__(self, k=1):
        self.k = k
        self.p = 2

    def _minkowski_dist(self, v1, v2):
        return sum(abs(e1 - e2) ** self.p for e1, e2 in zip(v1, v2)) ** (1 / self.p)

    def fit(self, X, labels, p=2):
        self.p = p
        self.X = X
        self.Y = labels
        return self
    
    def predict(self, query, method='mean'):
        assert method in ['mean', 'mode']
        dist_label = []
        for i in range(len(self.X)):
           dist_label.append((self._minkowski_dist(query, self.X[i]), self.Y[i]))

        neighbours = [x[1] for x in sorted(dist_label, key = lambda x: x[0])[:self.k]]
            
        if method=='mean':
            return np.average(neighbours, axis=0)
        elif method=='mode':
            count = Counter(neighbours)
            return max(count, key=count.get)

```

## Utility functions

```python
def plot_loss_acc(history):
    """Plot training and (optionally) validation loss and accuracy"""

    loss = history.history['loss']
    epochs = range(1, len(loss) + 1)

    plt.figure(figsize=(10, 10))

    plt.subplot(2, 1, 1)
    plt.plot(epochs, loss, '.--', label='Training loss')
    final_loss = loss[-1]
    title = 'Training loss: {:.4f}'.format(final_loss)
    plt.ylabel('Loss')
    if 'val_loss' in history.history:
        val_loss = history.history['val_loss']
        plt.plot(epochs, val_loss, 'o-', label='Validation loss')
        final_val_loss = val_loss[-1]
        title += ', Validation loss: {:.4f}'.format(final_val_loss)
    plt.title(title)
    plt.legend()

    acc = history.history['acc']

    plt.subplot(2, 1, 2)
    plt.plot(epochs, acc, '.--', label='Training acc')
    final_acc = acc[-1]
    title = 'Training accuracy: {:.2f}%'.format(final_acc * 100)
    plt.xlabel('Epochs')
    plt.ylabel('Accuracy')
    if 'val_acc' in history.history:
        val_acc = history.history['val_acc']
        plt.plot(epochs, val_acc, 'o-', label='Validation acc')
        final_val_acc = val_acc[-1]
        title += ', Validation accuracy: {:.2f}%'.format(final_val_acc * 100)
    plt.title(title)
    plt.legend()
```

## Step 1: Loading the data

### Question

* Load the Breast Cancer dataset included with scikit-learn. Store it in a variable named `dataset`. 
* Display 10 random samples with feature names, target and class.

<!-- grade_id: cell-438cd8e66cdd8522, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
dataset = load_breast_cancer()

df = pd.DataFrame(dataset['data'], columns=dataset['feature_names'])
#df['feature'] = dataset['data']
#dataset['feature_names']
df['target'] = dataset['target']
#df['class'] = dataset['target_names']
df['class'] = list(map(lambda x: 'malignant' if x == 1 else 'benign', dataset['target']))
df
```

Output:
```
mean radius  mean texture  mean perimeter  mean area  mean smoothness  \
0          17.99         10.38          122.80     1001.0          0.11840   
1          20.57         17.77          132.90     1326.0          0.08474   
2          19.69         21.25          130.00     1203.0          0.10960   
3          11.42         20.38           77.58      386.1          0.14250   
4          20.29         14.34          135.10     1297.0          0.10030   
..           ...           ...             ...        ...              ...   
564        21.56         22.39          142.00     1479.0          0.11100   
565        20.13         28.25          131.20     1261.0          0.09780   
566        16.60         28.08          108.30      858.1          0.08455   
567        20.60         29.33          140.10     1265.0          0.11780   
568         7.76         24.54           47.92      181.0          0.05263   

     mean compactness  mean concavity  mean concave points  mean sym
```

## Step 2: Using a simple model

```python
# Use scikit-learn's builtin logistic regression classifier to obtain a pretty good accuracy
model = LogisticRegression()
model.fit(dataset.data, dataset.target)
accuracy = model.score(dataset.data, dataset.target)
print('Accuracy: {:.2f}%'.format(accuracy * 100))
```

Output:
```
Accuracy: 94.55%
/home/roy/.conda/envs/MachineLearning/lib/python3.8/site-packages/sklearn/linear_model/_logistic.py:762: ConvergenceWarning: lbfgs failed to converge (status=1):
STOP: TOTAL NO. of ITERATIONS REACHED LIMIT.

Increase the number of iterations (max_iter) or scale the data as shown in:
    https://scikit-learn.org/stable/modules/preprocessing.html
Please also refer to the documentation for alternative solver options:
    https://scikit-learn.org/stable/modules/linear_model.html#logistic-regression
  n_iter_i = _check_optimize_result(
```

```python
def evaluate(y_hat, y_true):
    points = 0
    for pred, true in zip(y_hat, y_true):
        if pred == true:
            points += 1
    return points / len(y_hat)
```

```python
from sklearn.model_selection import train_test_split
X_train, X_test, y_train, y_test = train_test_split(dataset.data, dataset.target, 
                                                    stratify = dataset.target, test_size = 0.2,
                                                    random_state=4)

model = KNN(k = 3)
model.fit(X_train, y_train)
outcomes = []
for i in range(len(X_test)):
    outcomes.append(model.predict(X_test[i], method='mode'))
accuracy = evaluate(outcomes, y_test)
print('Accuracy: {:.2f}%'.format(accuracy * 100))
```

Output:
```
Accuracy: 94.74%
```

## Step 3: Training a neural network model

### Question

Train a model on the data to obtain a training accuracy > 85%. Store the training history in a variable named `history`.

<!-- grade_id: cell-400da8c44b9cae45, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

```python
plot_loss_acc(history)
```

<!-- grade_id: cell-c04f242e72fc4600, points: 1, grade: true, solution: false -->
```python
# Retrieve final accuracy
final_acc = history.history['acc'][-1]
# Assert final accuracy
assert_true(final_acc > 0.85)
```

## Step 4: Beating the simple model

### Question

Optimize training to beat the simple model by achieving a training accuracy > 96%.

<!-- grade_id: cell-b288a6495025292c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

```python
# Plot training history
plot_loss_acc(history)
```

<!-- grade_id: cell-cbe1f1188753f1ed, points: 2, grade: true, solution: false -->
```python
# Retrieve final accuracy
final_acc = history.history['acc'][-1]
# Assert final accuracy
assert_true(final_acc > 0.96)
```

```python

```